# HW1: Linear Regression

(Cập nhật lần cuối: 09/02/2025)

Họ tên: Vy Quốc Huy

MSSV: 22120142

---

## Nắm cách làm bài và nộp bài

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để phần nào giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.

*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>.  Bạn có thể thảo luận ý tưởng với bạn khác cũng như tham khảo các nguồn trên mạng, nhưng sau cùng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font> (khi tham khảo các nguồn trên mạng thì bạn cần ghi rõ nguồn trong bài làm, và đương nhiên là bạn cũng không được phép đưa code và bài làm cho bạn khác xem). <font color=red>Nếu vi phạm những điều này thì có thể bạn sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart Kernel & Run All Cells`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart Kernel & Run All Cells` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (ví dụ, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `HW1.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại với định dạng nén là .zip (chứ không được là .rar hay các định dạng khác) và nộp ở link trên moodle. \
<font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài này.</font>

## Kiểm tra môi trường code

```python
import sys
sys.executable
```

Output:
```
'/home/vyquochuy/miniconda3/envs/ml-env/bin/python'
```

Bạn nên thấy kết quả in ra là đường dẫn đến file chạy python của môi trường "ml-env" mà mình đã hướng dẫn bạn cài đặt ở HW0.

## Nắm bài toán cần giải quyết của bài tập này

Cho dữ liệu quan sát được (dữ liệu huấn luyện): 
$$\{(\textbf{x}^{(1)}, y^{(1)}), ..., (\textbf{x}^{(N_{train})}, y^{(N_{train})})\}$$

Trong đó:

- $\textbf{x}^{(n)} \in \mathbb{R}^{d}$ là véc-tơ đầu vào gồm $d$ phần tử (trong bài này, bạn sẽ làm với dữ liệu giả lập, chỉ đơn giản là những con số chứ không có ý nghĩa cụ thể; dữ liệu này có $d=2$, nhưng code của bạn vẫn nên chạy được với trường hợp $d \ge 1$ bất kỳ)
- $y^{(n)} \in \mathbb{R}$ là đầu ra tương ứng

Nhiệm vụ ở đây là tìm ra một mô hình (trong bài này: mô hình Linear Regression) từ dữ liệu này sao cho mô hình này có thể nhận một véc-tơ đầu vào *mới* ($\in \mathbb{R}^{d}$) và dự đoán đầu ra tương ứng một cách *chính xác* (việc dự đoán chính xác một cách tuyệt đối là rất khó, ở đây ta chỉ hướng đến là càng chính xác thì càng tốt).

## Import các thư viện cần thiết

```python
import numpy as np
import matplotlib.pyplot as plt
import itertools # Thư viện built-in của Python, sẽ được nói rõ hơn khi đụng đến ở cuối HW
plt.style.use('ggplot') # Để hình vẽ đẹp hơn một xíu ...
```

Lưu ý: khi tính toán với mảng Numpy, bạn nên dùng các toán-tử/hàm/phương-thức mà Numpy đã cung cấp sẵn. Các toán-tử/hàm/phương-thức này làm trên nguyên mảng và ở bên dưới đã được tối ưu hóa; do đó, code sẽ ngắn gọn và chạy nhanh. Nếu bạn dùng vòng lặp và làm với từng phần tử của mảng Numpy thì code sẽ dài và chạy chậm $\to$ bạn sẽ bị trừ điểm.

## Đọc dữ liệu

Trong bài này, bạn sẽ làm với dữ liệu do Thầy phát sinh ra bằng một hàm bí ẩn ;-). Dữ liệu ở đây chỉ đơn giản là những con số, chứ không có ý nghĩa cụ thể. Véc-tơ đầu vào gồm 2 phần tử, tức $d=2$. File "train.txt" chứa dữ liệu quan sát được (dữ liệu huấn luyện). File "new.txt" chứa dữ liệu mới, ngoài dữ liệu quan sát được.

Ở dưới, mình đã viết cho bạn hàm đọc dữ liệu từ file, và gọi hàm này để: (i) đọc dữ liệu từ file "train.txt" và lưu vào 2 mảng là `train_X` và `train_y`, (ii) đọc dữ liệu từ file "new.txt" và lưu vào 2 mảng là `new_X` và `new_y`.

```python
def read_data(data_file):
    data = np.loadtxt(data_file)
    X = data[:, 0:-1]
    y = data[:, -1]
    return X, y

train_X, train_y = read_data('train.txt')
new_X, new_y = read_data('new.txt')
print(f'Shape of train_X: {train_X.shape}, shape of train_y: {train_y.shape}')
print(f'Shape of new_X:   {new_X.shape}, shape of new_y:   {new_y.shape}')
```

Output:
```
Shape of train_X: (1000, 2), shape of train_y: (1000,)
Shape of new_X:   (1000, 2), shape of new_y:   (1000,)
```

## Khám phá dữ liệu quan sát được

Ta hãy xem thử min và max của dữ liệu quan sát được.

```python
feature1 = train_X[:, 0]
feature2 = train_X[:, 1]
print(f'Min of train_X[:, 0]: {feature1.min()}, max of train_X[:, 0]: {feature1.max()}')
print(f'Min of train_X[:, 1]: {feature2.min()}, max of train_X[:, 1]: {feature2.max()}')
print(f'Min of train_y:       {train_y.min()} , max of train_y:       {train_y.max()}')
```

Output:
```
Min of train_X[:, 0]: -0.998672, max of train_X[:, 0]: 0.999617
Min of train_X[:, 1]: -0.999265, max of train_X[:, 1]: 0.998556
Min of train_y:       -0.51948 , max of train_y:       0.944159
```

Như vậy, miền giá trị của mỗi cột trong `train_X` là từ xấp xỉ -1 đến xấp xỉ 1, còn miền giá trị của `train_y` là từ xấp xỉ -0.5 đến xấp xỉ 1.

## Tiền xử lý dữ liệu quan sát được

Để có thể tính ra bộ trọng số của mô hình Linear Regression (theo công thức ở trang 13 ở slide lý thuyết "lect03...") thì trước tiên ta cần thêm một cột với các giá trị đều bằng 1 vào đầu mảng `train_X`. Mảng `train_X` sau khi được thêm cột 1 chính là ma trận $\textbf{X}$ trong slide lý thuyết, và ta sẽ sẵn sàng để tính ra bộ trọng số của Linear Regression.

***Nhiệm vụ của bạn (2đ):***

- Viết hàm `add_ones`. Hàm này có tham số đầu vào là mảng `X` có shape là `(N, d)` với `N` là số lượng véc-tơ đầu vào và `d` là số lượng phần tử của mỗi véc-tơ đầu vào. Hàm này sẽ trả về mảng `X` sau khi đã được thêm cột 1 ở đầu và mảng này có shape là `(N, d+1)`.
- Gọi hàm `add_ones` với tham số đầu vào là `train_X` và lưu kết quả vào biến `train_Z`.

<!-- grade_id: cell-623796af0978e3ec, points: 0, grade: false, solution: true -->
```python
def add_ones(X):
    # YOUR CODE HERE
    N = X.shape[0]
    ones = np.ones((N, 1))
    return np.hstack((ones, X))

# Gọi hàm add_ones để tiền xử lý train_X
# YOUR CODE HERE
train_Z = add_ones(train_X)
```

<!-- grade_id: cell-bfb56c020dca8c78, points: 2, grade: true, solution: false -->
```python
# TEST
assert train_Z.shape == (1000, 3)
assert str(train_Z[0][0].round(4)) == '1.0'
assert str(train_Z[0][1].round(4)) == '0.0976'
assert str(train_Z[1][0].round(4)) == '1.0'
assert str(train_Z[1][1].round(4)) == '0.2055'
assert str(train_Z[2][0].round(4)) == '1.0'
assert str(train_Z[2][1].round(4)) == '-0.1527'
```

## Tìm mô hình Linear Regression từ dữ liệu quan sát được

Sau khi đã có được mảng `train_Z` (là `train_X` đã được thêm cột 1) thì ta có thể dễ dàng tìm được bộ trọng số của Linear Regression (theo công thức ở trang 13 ở slide lý thuyết "lect03...").

***Nhiệm vụ của bạn (2đ):***

- Viết hàm `train_linreg`. Hàm này có tham số đầu vào là mảng `X` và mảng `y`. Mảng `X` chứa các véc-tơ đầu vào của dữ liệu quan sát được, mảng này đã được thêm cột 1 ở đầu, có shape là `(N, d+1)` với `N` là số lượng véc-tơ đầu vào của dữ liệu quan sát được và `d` là số lượng phần tử của mỗi véc-tơ đầu vào khi chưa thêm 1 ở đầu. Mảng `y` chứa các đầu ra tương ứng, có shape là `(N,`). Hàm này trả về mảng chứa các trọng số của mô hình Linear Regression tìm được, có shape là `(d+1,)`. \
Theo công thức ở trang 13 ở slide lý thuyết "lect03..." thì ta sẽ cần tính nghịch đảo của `X` chuyển vị nhân `X`. Trong đa số trường hợp thì `X` chuyển vị nhân `X` sẽ khả nghịch, nhưng đôi khi xui quá cũng có thể xảy ra trường hợp không khả nghịch. Để chắc ăn, bạn hãy dùng hàm `np.linalg.pinv` (thay vì `np.linalg.inv`) để tích nghịch đảo; trong trường hợp không khả nghịch, hàm này vẫn sẽ tính ra cho bạn một kết quả xấp xỉ.  
- Gọi hàm `train_linreg` với tham số đầu vào là `train_Z` và `train_y`, lưu bộ trọng số tìm được vào biến `w`.

<!-- grade_id: cell-fa693a19ef4c6057, points: 0, grade: false, solution: true -->
```python
def train_linreg(X, y):
    # YOUR CODE HERE
    # Tính w = (Z^T * Z)^(-1) * Z^T * y
    ZT = X.T                # Lấy ma trận chuyển vị của X
    ZTZ = ZT.dot(X)         # Nhân Z^T với Z 
    ZTZ_inv = np.linalg.pinv(ZTZ) # Tìm ma trận nghịch đảo của ZTZ
    ZTy = ZT.dot(y)               
    return ZTZ_inv.dot(ZTy) # Trả về w

# Gọi hàm train_linreg để tìm bộ trọng số của Linear Regression từ train_Z và train_y
# YOUR CODE HERE
w = train_linreg(train_Z, train_y)
```

<!-- grade_id: cell-c87b909528b6bde1, points: 2, grade: true, solution: false -->
```python
# TEST
assert w.shape == (3,)
assert str(w[0].round(4)) == '-0.1484'
assert str(w[1].round(4)) == '0.0131'
assert str(w[2].round(4)) == '0.018'
```

## Dùng mô hình Linear Regression tìm được để dự đoán với dữ liệu quan sát được và đánh giá kết quả

Sau khi tìm được bộ trọng số của Linear Regression thì ta sẽ thử dùng bộ trọng số này để dự đoán với các véc-tơ đầu vào của dữ liệu quan sát được và so sánh với các đầu ra đúng để xem sai lệch là bao nhiêu.

***Nhiệm vụ 1 của bạn (2đ):***

- Viết hàm `compute_linreg_output`. Hàm này có tham số đầu vào là mảng `w` và mảng `X`. Mảng `w` chứa bộ trọng số của Linear Regresssion, có shape là `(d+1,)` với `d` là số lượng phần tử của véc-tơ đầu vào khi chưa thêm 1 ở đầu. Mảng `X` chứa các véc-tơ đầu vào cần dự đoán, mảng này đã được thêm cột 1 ở đầu, có shape là `(N, d+1)` với `N` là số lượng véc-tơ đầu vào. Hàm này trả về mảng chứa các đầu ra dự đoán, có shape là `(N,)`.
- Gọi hàm `compute_linreg_output` với tham số đầu vào là mảng `w` và mảng `train_Z`, lưu kết quả vào biến `predicted_train_y`.

<!-- grade_id: cell-4871d7e329ece947, points: 0, grade: false, solution: true -->
```python
def compute_linreg_output(w, X):
    # YOUR CODE HERE
    return X.dot(w) # Trả về giá trị dự đoán

# Gọi hàm compute_linreg_output để dự đoán với train_Z
# YOUR CODE HERE
predicted_train_y = compute_linreg_output(w, train_Z)
```

<!-- grade_id: cell-a2b5aec889837078, points: 2, grade: true, solution: false -->
```python
# TEST
assert predicted_train_y.shape == (1000,)
assert str(predicted_train_y[0].round(4)) == '-0.1394'
assert str(predicted_train_y[1].round(4)) == '-0.1441'
assert str(predicted_train_y[2].round(4)) == '-0.1451'
```

***Nhiệm vụ 2 của bạn (2đ):***

- Viết hàm `compute_mae`. Hàm này có tham số đầu vào là mảng `predicted_y` và mảng `y`. Mảng `predicted_y` chứa các đầu ra dự đoán, có shape là `(N,)`. Mảng `y` chứa các đầu ra đúng, cũng có shape là `(N,)`. Hàm này trả về một số thực cho biết độ lỗi MAE (Mean Absolute Error) giữa 2 mảng `predicted_y` và `y`. Công thức tính MAE thì tương tự như MSE (Mean Square Error) nhưng thay bình phương bằng trị tuyệt đối. Ở đây, ta dùng MAE vì MAE dễ cảm nhận hơn so với MSE.
- Gọi hàm `compute_mae` với tham số đầu vào là mảng `predicted_train_y` và `train_y`, lưu kết quả vào biến `train_mae`.

<!-- grade_id: cell-0db101b4f37e16a7, points: 0, grade: false, solution: true -->
```python
def compute_mae(predicted_y, y):
    # YOUR CODE HERE
    return np.mean(np.abs(predicted_y - y)) # Trả về giá trị MAE

# Gọi hàm compute_mae để tính độ lỗi giữa predicted_train_y và train_y
# YOUR CODE HERE
train_mae = compute_mae(predicted_train_y, train_y)
```

<!-- grade_id: cell-18aa7cf8fbcb9247, points: 2, grade: true, solution: false -->
```python
# TEST
assert str(np.round(train_mae, 4)) == '0.2037'
```

Hmm .... Với `train_y` có miền giá trị từ xấp xỉ -0.5 đến xấp xỉ 1 thì có thể nói độ lỗi `train_mae` này là khá lớn.

Nhưng cái ta quan tâm sau cùng là độ lỗi của mô hình Linear Regression tìm được trên dữ liệu mới ngoài dữ liệu quan sát được. Biết đâu ... Ở phần kế, ta sẽ xem thử vận may của ta thế nào ;-).

## Tiền xử lý dữ liệu mới, dùng mô hình Linear Regression tìm được để dự đoán với dữ liệu mới và đánh giá kết quả

```python
new_Z = add_ones(new_X)
predicted_new_y = compute_linreg_output(w, new_Z)
new_mae = compute_mae(predicted_new_y, new_y)
assert str(np.round(new_mae, 4)) == '0.2065'
```

Vận may đã không xảy ra :-(. Độ lỗi `new_mae` vẫn lớn không kém gì `train_mae`, thậm chí `new_mae` còn lớn hơn `train_mae` một xíu. 

Nếu muốn `new_mae` nhỏ thì ta phải suy nghĩ cách cải tiến, chứ không thể dựa vào vận may.

## Đưa ra ý tưởng cải tiến

Ưu điểm của mô hình Linear Regression là đơn giản, huấn luyện nhanh. Nhưng đơn giản cũng là nhược điểm của mô hình Linear Regression. Qua kết quả `train_mae` ở trên, ta thấy mô hình Linear Regression không khớp tốt vào dữ liệu quan sát được, dẫn tới độ lỗi lớn trên dữ liệu quan sát được. Nếu trên dữ liệu quan sát được mà có độ lỗi lớn thì trên dữ liệu mới thường cũng sẽ có độ lỗi lớn.

Nếu ta có thể cải thiện độ lỗi của mô hình Linear Regression trên dữ liệu quan sát được thì độ lỗi trên dữ liệu mới có khả năng sẽ được cải thiện theo. Một cách để cải thiện độ lỗi của mô hình Linear Regression trên dữ liệu quan sát được là ta sẽ suy nghĩ để thiết kế ra véc-tơ đầu vào $\textbf{z}$ thay cho véc-tơ đầu vào $\textbf{x}$ sao cho $\textbf{z}$ sẽ có mối quan hệ tuyến tính (hoặc xấp xỉ tuyến tính) với $y$ trên dữ liệu quan sát được; sau đó thì ta có thể áp dụng Linear Regression một cách bình thường nhưng thay $\textbf{x}$ bằng $\textbf{z}$. 

Với dữ liệu mà ta đang làm, véc-tơ đầu vào $\textbf{x}$ có $d=2$, tức gồm 2 phần tử là $x_1$ và $x_2$. Ta hình dung: trong hệ trục tọa độ 3 chiều mà trục đứng là $y$ và 2 trục còn lại là $x_1$ và $x_2$, sẽ có các điểm dữ liệu ứng với dữ liệu quan sát được và mô hình Linear Regression $y = w_0 + w_1x_1 + w_2x_2$ là một *mặt phẳng* trong hệ trục tọa độ này; việc tìm mô hình Linear Regression mà bạn đã làm ở trên chính là tìm mặt phẳng khớp nhất với các điểm dữ liệu quan sát được. Kết quả `train_mae` cho thấy độ lỗi lớn, nghĩa là việc dùng *mặt phẳng* không khớp tốt vào các điểm dữ liệu quan sát được. Vậy để có thể khớp tốt hơn, ta sẽ dùng *mặt cong*. Cụ thể, ta có thể thiết kế véc-tơ đầu vào $\textbf{z}$ là véc-tơ đầu vào $\textbf{x}$ ($\textbf{x}$ gồm 2 phần tử là $x_1$ và $x_2$) được bổ sung thêm các phần tử bậc 2 nữa là $x_1^2, x_1x_2, x_2^2$; lúc này mô hình Linear Regression sẽ có dạng mặt cong $y = w_0 + w_1z_1 + w_2z_2 + w_3z_3 + w_4z_4 + w_5z_5$ với $z_1, z_2, z_3, z_4, z_5$ là các thành phần của véc-tơ đầu vào $\textbf{z}$, và $z_1 = x_1, z_2 = x_2, z_3 = x_1^2, z_4 = x_1x_2, z_5 = x_2^2$.

Tiếp theo, ta sẽ thử cài đặt ý tưởng này!

## Tiền xử lý dữ liệu quan sát được

Từ mảng `train_X`, ta sẽ tạo ra mảng `train_Z` bằng cách: thêm các cột bậc 2 vào cuối `train_X`, rồi thêm cột 1 vào đầu `train_X`. 

***Nhiệm vụ của bạn (2đ):***

- Viết hàm `add_squares`. Hàm này có tham số đầu vào là mảng `X`. Mảng `X` chứa các véc-tơ đầu vào (chưa thêm 1 ở đầu), có shape là `(N, d)` với `N` là số lượng véc-tơ đầu vào và `d` là số lượng phần tử của mỗi véc-tơ đầu vào. Hàm này trả về mảng `X` sau khi đã được thêm các cột vào cuối mảng, và mảng này có shape là `(N, d + (d+2-1)!/2!/(d-1)!)` (ví dụ, với `X` ban đầu có `d=2` cột là `x_1` và `x_2` thì mảng kết quả sẽ có 2 + (2+2-1)!/2!/(2-1)! = 5 cột: `x_1`, `x_2`, `x_1` nhân `x_1`, `x_1` nhân `x_2`, `x_2` nhân `x_2`). \
Trong hàm `add_squares`, bạn sẽ dùng câu lệnh\
`list(itertools.combinations_with_replacement(np.arange(X.shape[1]), 2))` \
để phát sinh ra các tổ hợp bậc 2 có thể có từ các cột của `X`. Ví dụ, với `X` gồm 2 cột là cột 0 và cột 1 thì câu lệnh trên sẽ ra list kết quả là `[(0, 0), (0, 1), (1, 1)]`. Trong hàm `add_squares`, bạn được phép dùng một vòng lặp để duyệt qua các phần tử trong list kết quả ở trên; đây cũng là chỗ duy nhất trong HW1 mà bạn được phép dùng vòng lặp.
- Gọi hàm `add_squares` với tham số đầu vào là `train_X`, rồi gọi tiếp hàm `add_ones` với tham số đầu vào là kết quả sau khi gọi hàm `add_squares`. Lưu kết quả sau cùng vào biến `train_Z`.

<!-- grade_id: cell-3ad36330bbe4eeba, points: 0, grade: false, solution: true -->
```python
def add_squares(X):
    # YOUR CODE HERE
    N, d = X.shape
    combos = list(itertools.combinations_with_replacement(np.arange(X.shape[1]), 2))  # Các cặp chỉ số (i, j)
    
    poly_features = []  # Danh sách các cột mới
    for i, j in combos:
        new_col = (X[:, i] * X[:, j]).reshape(-1, 1)  # Nhân từng cặp cột và reshape lại thành cột
        poly_features.append(new_col)
    
    poly_features = np.hstack(poly_features)  # Ghép tất cả các cột lại
    return np.hstack((X, poly_features))  # Nối X gốc với phần bậc 2

# Gọi hàm add_squares rồi add_ones để tiền xử lý train_X
# YOUR CODE HERE
train_Z = add_ones(add_squares(train_X))
```

<!-- grade_id: cell-cb13aaae72dc9884, points: 2, grade: true, solution: false -->
```python
# TEST 1 (d = 2)
assert train_Z.shape == (1000, 6)
assert str(train_Z[0][0].round(4)) == '1.0'
assert str(train_Z[0][1].round(4)) == '0.0976'
assert str(train_Z[0][2].round(4)) == '0.4304'
assert str(train_Z[0][3].round(4)) == '0.0095'
assert str(train_Z[0][4].round(4)) == '0.042'
assert str(train_Z[0][5].round(4)) == '0.1852'

# TEST 2 (d = 3)
X = np.arange(6).reshape(2, 3)
Z = add_squares(X)
assert Z.shape == (2, 9)
assert Z[0][0] == 0
assert Z[0][1] == 1
assert Z[0][2] == 2
assert Z[0][3] == 0
assert Z[0][4] == 0
assert Z[0][5] == 0
assert Z[0][6] == 1
assert Z[0][7] == 2
assert Z[0][8] == 4
```

## Tìm mô hình Linear Regression từ dữ liệu quan sát được

```python
w = train_linreg(train_Z, train_y)
```

## Dùng mô hình Linear Regression tìm được để dự đoán với dữ liệu quan sát được và đánh giá kết quả

```python
predicted_train_y = compute_linreg_output(w, train_Z)
train_mae = compute_mae(predicted_train_y, train_y)
assert str(np.round(train_mae, 4)) == '0.0078'
```

Woohoo!

Khoan khoan, cái ta quan tâm là độ lỗi trên dữ liệu mới ngoài dữ liệu quan sát được.

## Tiền xử lý dữ liệu mới, dùng mô hình Linear Regression tìm được để dự đoán với dữ liệu mới và đánh giá kết quả

```python
new_Z = add_ones(add_squares(new_X))
predicted_new_y = compute_linreg_output(w, new_Z)
new_mae = compute_mae(predicted_new_y, new_y)
assert str(np.round(new_mae, 4)) == '0.008'
```

Woohoo!

Như vậy là độ lỗi MAE trên cả dữ liệu quan sát được lẫn dữ liệu mới đều giảm xuống một cách đáng kể. Happy ending :-).